---
title: 'Intro and Setup'
description: 'The first steps needed to successfully get your workstation up and running with Python'
---

## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 4413628931   
```





`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```


`@sct`

```{python}

```


